# Enigma Machine
## Simulation


<!-- grade_id: cell-10268453ab40f7dd, points: 0, grade: false, solution: true -->
```python
from enigma import *
from decryptor import *
```

### Plug Leads
Plug leads are the outermost step of the process.

Plug leads bind two characters together. 

Every time character 1 is introduced as an input, the output of character 2 is passed on to the rest of the process.

Since the machine is reflected by the reflector, the plug scramble letters twice; once after collecting an input from the user and again just before a character is outputted by the machine.

<!-- grade_id: cell-823cd8cc358d20ff, points: 5, grade: true, solution: false -->
```python
lead = PlugLead("AG")
assert(lead.encode("A") == "G")
assert(lead.encode("D") == "D")

lead = PlugLead("DA")
assert(lead.encode("A") == "D")
assert(lead.encode("D") == "A")

# Unmapped characters should return unchanged
assert(lead.encode("R") == "R")
```

### Plugboard
The plugboard simply houses all of our plug leads.
The original enigma machine had a plug board that could house up to 10 leads. Operators had a calendar that would dictate the composition of plug leads to use for every given day.

<!-- grade_id: cell-aa9f93cfa0811bf2, points: 5, grade: true, solution: false -->
```python
plugboard = Plugboard()

plugboard.add(PlugLead("SZ"))
plugboard.add(PlugLead("GT"))
plugboard.add(PlugLead("DV"))
plugboard.add(PlugLead("KU"))

assert(plugboard.encode("K") == "U")
assert(plugboard.encode("A") == "A")
```

### Rotors
The enigma machine uses a composition of different rotors that are responsible for scrambling a character before passing it forward. The outermost rotor turns every time it is used. After a full rotation back to its starting point, the second rotor turns once. After the second rotor makes a full rotation, the third rotor will turn, and so on for n rotors.

Each rotor can be chosen from a box containing seven possible wiring patterns. There are two rotors labelled `Beta` and `Gamma`. Then there are five rotors labelled with Roman numerals which do rotate: `I, II, III, IV, V`. Finally, the machine has three different types of reflector wiring patterns, labelled `A, B, C`.

Wiring patterns from wikipedia: [this page](https://en.wikipedia.org/wiki/Enigma_rotor_details).
<table><thead>
<tr><th></th><th colspan="26"><center>Mapping from letter</center></th></tr>
<tr><th style="text-align:left">Label</th><th>A</th><th>B</th><th>C</th><th>D</th><th>E</th><th>F</th><th>G</th><th>H</th><th>I</th><th>J</th><th>K</th><th>L</th><th>M</th><th>N</th><th>O</th><th>P</th><th>Q</th><th>R</th><th>S</th><th>T</th><th>U</th><th>V</th><th>W</th><th>X</th><th>Y</th><th>Z</th></tr></thead><tbody>
<tr><th style="text-align:left">Beta</th><td>L</td><td>E</td><td>Y</td><td>J</td><td>V</td><td>C</td><td>N</td><td>I</td><td>X</td><td>W</td><td>P</td><td>B</td><td>Q</td><td>M</td><td>D</td><td>R</td><td>T</td><td>A</td><td>K</td><td>Z</td><td>G</td><td>F</td><td>U</td><td>H</td><td>O</td><td>S</td></tr>
<tr><th style="text-align:left">Gamma</th><td>F</td><td>S</td><td>O</td><td>K</td><td>A</td><td>N</td><td>U</td><td>E</td><td>R</td><td>H</td><td>M</td><td>B</td><td>T</td><td>I</td><td>Y</td><td>C</td><td>W</td><td>L</td><td>Q</td><td>P</td><td>Z</td><td>X</td><td>V</td><td>G</td><td>J</td><td>D</td></tr>
<tr><th style="text-align:left">I</th><td>E</td><td>K</td><td>M</td><td>F</td><td>L</td><td>G</td><td>D</td><td>Q</td><td>V</td><td>Z</td><td>N</td><td>T</td><td>O</td><td>W</td><td>Y</td><td>H</td><td>X</td><td>U</td><td>S</td><td>P</td><td>A</td><td>I</td><td>B</td><td>R</td><td>C</td><td>J</td></tr>
<tr><th style="text-align:left">II</th><td>A</td><td>J</td><td>D</td><td>K</td><td>S</td><td>I</td><td>R</td><td>U</td><td>X</td><td>B</td><td>L</td><td>H</td><td>W</td><td>T</td><td>M</td><td>C</td><td>Q</td><td>G</td><td>Z</td><td>N</td><td>P</td><td>Y</td><td>F</td><td>V</td><td>O</td><td>E</td></tr>
<tr><th style="text-align:left">III</th><td>B</td><td>D</td><td>F</td><td>H</td><td>J</td><td>L</td><td>C</td><td>P</td><td>R</td><td>T</td><td>X</td><td>V</td><td>Z</td><td>N</td><td>Y</td><td>E</td><td>I</td><td>W</td><td>G</td><td>A</td><td>K</td><td>M</td><td>U</td><td>S</td><td>Q</td><td>O</td></tr>
<tr><th style="text-align:left">IV</th><td>E</td><td>S</td><td>O</td><td>V</td><td>P</td><td>Z</td><td>J</td><td>A</td><td>Y</td><td>Q</td><td>U</td><td>I</td><td>R</td><td>H</td><td>X</td><td>L</td><td>N</td><td>F</td><td>T</td><td>G</td><td>K</td><td>D</td><td>C</td><td>M</td><td>W</td><td>B</td></tr>
<tr><th style="text-align:left">V</th><td>V</td><td>Z</td><td>B</td><td>R</td><td>G</td><td>I</td><td>T</td><td>Y</td><td>U</td><td>P</td><td>S</td><td>D</td><td>N</td><td>H</td><td>L</td><td>X</td><td>A</td><td>W</td><td>M</td><td>J</td><td>Q</td><td>O</td><td>F</td><td>E</td><td>C</td><td>K</td></tr>
<tr><th style="text-align:left">A</th><td>E</td><td>J</td><td>M</td><td>Z</td><td>A</td><td>L</td><td>Y</td><td>X</td><td>V</td><td>B</td><td>W</td><td>F</td><td>C</td><td>R</td><td>Q</td><td>U</td><td>O</td><td>N</td><td>T</td><td>S</td><td>P</td><td>I</td><td>K</td><td>H</td><td>G</td><td>D</td></tr>
<tr><th style="text-align:left">B</th><td>Y</td><td>R</td><td>U</td><td>H</td><td>Q</td><td>S</td><td>L</td><td>D</td><td>P</td><td>X</td><td>N</td><td>G</td><td>O</td><td>K</td><td>M</td><td>I</td><td>E</td><td>B</td><td>F</td><td>Z</td><td>C</td><td>W</td><td>V</td><td>J</td><td>A</td><td>T</td></tr>
<tr><th style="text-align:left">C</th><td>F</td><td>V</td><td>P</td><td>J</td><td>I</td><td>A</td><td>O</td><td>Y</td><td>E</td><td>D</td><td>R</td><td>Z</td><td>X</td><td>W</td><td>G</td><td>C</td><td>T</td><td>K</td><td>U</td><td>Q</td><td>S</td><td>B</td><td>N</td><td>M</td><td>H</td><td>L</td></tr>
</tbody></table>

### Enigma Machine Demonstration


<!-- grade_id: cell-993ccf0331dfd6eb, points: 10, grade: true, solution: true -->
```python
    plugboard = Plugboard()

    plugboard.add(PlugLead("HL"))
    plugboard.add(PlugLead("MO"))
    plugboard.add(PlugLead("AJ"))
    plugboard.add(PlugLead("CX"))
    plugboard.add(PlugLead("BZ"))
    plugboard.add(PlugLead("SR"))
    plugboard.add(PlugLead("NI"))
    plugboard.add(PlugLead("YW"))
    plugboard.add(PlugLead("DG"))
    plugboard.add(PlugLead("PK"))

    board = Rotorboard(["I", "II", "III"], "B",[1,1,1], "AAZ")
    
    enigma = EnigmaMachine(board, plugboard)
    input1 = "HELLOWORLD"
    output1 = enigma.EncodeMessage(input1)
    print(f"{input1} translates to: {output1}")
   
    
    plugboard2 = Plugboard()
    plugboard2.add(PlugLead("PC"))
    plugboard2.add(PlugLead("XZ"))
    plugboard2.add(PlugLead("FM"))
    plugboard2.add(PlugLead("QA"))
    plugboard2.add(PlugLead("ST"))
    plugboard2.add(PlugLead("NB"))
    plugboard2.add(PlugLead("HY"))
    plugboard2.add(PlugLead("OR"))
    plugboard2.add(PlugLead("EV"))
    plugboard2.add(PlugLead("IU"))
    
    board2 = Rotorboard(["IV", "V", "Beta", "I"], "A",[18,24,3,5], "EZGP")
    
    enigma2 = EnigmaMachine(board2, plugboard2)
    input2 = "MESSAGE"
    output2 = enigma2.EncodeMessage(input2)
    print(f"{input2} translates to: {output2}")
    
    board2 = Rotorboard(["IV", "V", "Beta", "I"], "A",[18,24,3,5], "EZGP")
    enigma2 = EnigmaMachine(board2, plugboard2)
    print(f"It translates back to {enigma2.EncodeMessage(output2)}")
```

Output:
```
HELLOWORLD translates to: RFKTMBXVVW
MESSAGE translates to: YPNDYST
It translates back to MESSAGE
```

```python

```